Before you turn this problem in, make sure everything runs as expected. **Restart the kernel and run all cells** (in the menubar, select Kernel$\rightarrow$Restart & Run All).

Make sure you fill in any place that says `YOUR CODE HERE ???` or "YOUR ANSWER HERE". Code cells are autograded, and if the test cells don't run all the way to the line that  prints "Tests Passed (# points)", you will not recieve any point for that test.

# CSCI 3155: Assignment 8

Goals of this assignment include: 
- Developing a type system for references
- Exercises on type checking and type inference posted on moodle.

__Name__: Derek Thomas

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (25 points)

In this problem, we will extend our type system for handling references.  Here is a stripped down version of Lettuce that involves references.

$$\begin{array}{rcl}
\mathbf{Expr} & \rightarrow & Const(\mathbf{Double}) \\
& | & Ident(\mathbf{Identifier})\\
& | & Plus(\mathbf{Expr}, \mathbf{Expr})\\
& | & Div(\mathbf{Expr}, \mathbf{Expr}) \\
& | & Geq(\mathbf{Expr}, \mathbf{Expr}) \\
& | & And(\mathbf{Expr}, \mathbf{Expr}) \\
& | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr})\\
& | & Let(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}, \mathbf{Expr})\\
& | & FunDef(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}) \\
& | & FunCall(\mathbf{Expr}, \mathbf{Expr}) \\
& | & \color{blue}{NewRef(\mathbf{Expr})} \\
& | & \color{blue}{DeRef(\mathbf{Expr})} \\
& | & \color{blue}{AssignRef(\mathbf{Expr}, \mathbf{Expr})} \\
\end{array}$$

We have a type system 
$$\begin{array}{rcl}
\mathbf{Type} & \rightarrow & NumType & \leftarrow \text{written as } \ num \\
& | & BoolType & \leftarrow \text{written as } \ bool \\
& | & FunType(\mathbf{Type} , \mathbf{Type} ) & \leftarrow \text{written as } \ t_1 \Rightarrow t_2 \\
& | & \color{red}{RefType(\mathbf{Type})} & \leftarrow \text{written as } \ ref(t_1) \\
\end{array}$$


### A (5 points)
Fill out the missing type annotations labeled type1, type2, and type3 in the program below.

~~~
let x: type1 = 25 in 
    let y : type2 = NewRef( - x ) in 
     let f : type3 = NewRef( function (z: ref(num)) DeRef(z) - x >= 45 ) in 
          let z : ref(bool) = DeRef(f)(y)
~~~

<!-- grade_id: cell-312b0c7aebe54f39, points: 5, grade: true, solution: true -->
type1: NumType
type2: RefType(NumType)
type2: RefType(FunType(RefType(NumType), BoolType))

### B (10 points): Let us write some inference rules for working with ref types.
$$\newcommand\typeOf{\mathbf{typeOf}}$$
$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array}\;(\text{#3}) }$$

Recall from notes on "Types and Type Checking" that $\typeOf(\texttt{e}, \alpha)$ is the type of an 
expression $\texttt{e}$ under type environment $\alpha$. The type environment maps identifiers in the current scope to  their annotated types.


Let us write a rule for NewRef.

$$\semRule{ \typeOf(\texttt{e}, \alpha) = t,\; t \not= \mathbf{typeerror} }{\typeOf(\texttt{NewRef(e)}, \alpha) = ref(t) }{newref-ok}$$

It says that if $\texttt{e}$ receives type $t$ under type environment $\alpha$ and it is not a type error, then $\texttt{NewRef(e)}$ must receive the type $ref(t)$ under $\alpha$.



(i) Complete the missing terms for the rule for `DeRef` OK rule.
$$\semRule{\typeOf(\texttt{e}, \alpha) = ref(t)}{\typeOf(\texttt{DeRef(e)}, \alpha) = \color{red}{???_1}}{deref-ok}$$

(ii) Complete the missing terms for the rule for `DeRef` error.
$$\semRule{\typeOf(\texttt{e}, \alpha) = t,\ t \not= \color{red}{???_2} }{\typeOf(\texttt{DeRef(e)}, \alpha) = \mathbf{typeerror}}{deref-nok}$$

(iii) Complete the missing terms for `AssignRef` OK rule.
$$\semRule{\typeOf(\texttt{e1}, \alpha) = \color{red}{???_3}, \typeOf(\texttt{e2}, \alpha) = t}{\typeOf(\texttt{AssignRef(e1, e2)}, \alpha) =   \color{red}{???_4}}{assignref-ok}$$


<!-- grade_id: cell-e1cc8dd790c95fef, points: 10, grade: true, solution: true -->
1. t
2. ref(t)
3. ref(t)
4. t

### C (10 points): Write the type checker with references by filling in the missing code.

<!-- grade_id: cell-3b12181f57c0872f, points: 0, grade: false, solution: true -->
```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType(t1: Type, t2: Type) extends Type
// TODO: Write a new case class for RefType
// YOUR CODE HERE
case class RefType(t1: Type) extends Type

sealed trait Program
sealed trait Expr

case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Geq(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x: String, xType: Type, e1: Expr, e2: Expr) extends Expr
case class FunDef(id: String, idType: Type, e: Expr) extends Expr
case class FunCall(calledFun: Expr, argExpr: Expr) extends Expr
case class NewRef(e: Expr) extends Expr
case class DeRef(e: Expr) extends Expr
case class AssignRef(e1: Expr, e2: Expr) extends Expr
case class TopLevel(e: Expr) extends Program

def typeEquals(t1: Type, t2: Type): Boolean = t1 == t2
case class TypeErrorException(s: String) extends Exception

```

Output:
```
defined [32mtrait[39m [36mType[39m
defined [32mobject[39m [36mNumType[39m
defined [32mobject[39m [36mBoolType[39m
defined [32mclass[39m [36mFunType[39m
defined [32mclass[39m [36mRefType[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mNewRef[39m
defined [32mclass[39m [36mDeRef[39m
defined [32mclass[39m [36mAssignRef[39m
defined [32mclass[39m [36mTopLevel[39m
defined [32mfunction[39m [36mtypeEquals[39m
defined [32mclass[39m [36mTypeErrorException[39m
```

<!-- grade_id: cell-e8a060ea96ec6931, points: 0, grade: false, solution: true -->
```scala
def typeOf(e: Expr, alpha: Map[String, Type]): Type = {
    def checkType(opName: String, e1: Expr, t1: Type, e2: Expr, t2: Type, resType: Type): Type = {
        val t1hat = typeOf(e1, alpha)
        if (! typeEquals(t1hat, t1)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t1, obtained $t1hat")
        }
        
        val t2hat = typeOf(e2, alpha)
        if (! typeEquals(t2hat, t2)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t2, obtained $t2hat")
        }
        
        resType
    }
    
    e match {
        case Const(f) => NumType
        case Ident(s) => {if (alpha contains s)
                             alpha(s)
                          else 
                             throw TypeErrorException(s"Unknown identifier $s")}
        case Plus(e1, e2) =>  checkType("Plus", e1,  NumType, e2, NumType, NumType)
        case Minus(e1, e2) => checkType("Minus",e1,  NumType, e2, NumType, NumType)
        case Geq(e1, e2) => checkType("Geq", e1,  NumType, e2, NumType, BoolType)
        case IfThenElse(e, e1, e2) => {
            val t = typeOf(e, alpha)
            if (t == BoolType){
                val t1 = typeOf(e1, alpha)
                val t2 = typeOf(e2, alpha)
                if (typeEquals(t1, t2))
                    t1
                else 
                    throw TypeErrorException(s"If then else returns unequal types $t1 and $t2")
            } else {
                throw TypeErrorException(s"If then else condition expression not boolean $t")
            }
        }

        case Let(x, t, e1, e2) => {
            val t1 = typeOf(e1, alpha)
            if (typeEquals(t1, t)){
                val newAlpha = alpha + (x -> t)
                typeOf(e2, newAlpha)
            } else {
                throw TypeErrorException(s"Let binding has type $t whereas it is bound to expression of type $t1")
            }
        }

        case FunDef(x, t1, e) => {
            val newAlpha = alpha + (x -> t1)
            val t2 = typeOf(e, newAlpha)
            FunType(t1, t2)
        }

        case FunCall(e1, e2) => {
            val ftype = typeOf(e1, alpha)
            ftype match {
                case FunType(t1, t2) => {
                    val argType = typeOf(e2, alpha)
                    if (typeEquals(argType, t1)){
                        t2
                    } else {
                        throw TypeErrorException(s"Call to function with incompatible argument type. Expected $t1, obtained $argType")
                    }
                }
                case _ => { throw TypeErrorException(s"Call to function but with a non function type $ftype")}

            }
        }

        case NewRef(e) => {
            // YOUR CODE HERE
            val rtype = typeOf(e, alpha)
            rtype match {
                case NumType => {
                    RefType(typeOf(e, alpha))
                }
                case BoolType => {
                    RefType(typeOf(e, alpha))
                }
                case FunType(t1,t2) => {
                    RefType(typeOf(e, alpha))
                }
                case RefType(t1) => {
                    RefType(typeOf(e, alpha))
                }
                case _ => { throw TypeErrorException(s"Expected Type")}
            }
            
        }
        
        case AssignRef(e1, e2) => {
            // YOUR CODE HERE
            val rtype = typeOf(e1, alpha)
            rtype match {
                case RefType(t1) => {
                    typeOf(e2, alpha)
                }
                case _ => { throw TypeErrorException(s"Call to AssignRef but with a non reference type $rtype")}
            }
        }
        
        case DeRef(e) => {
            // YOUR CODE HERE
            val rtype = typeOf(e, alpha)
            rtype match {
                case RefType(t1) => {
                    t1
                }
                case _ => { throw TypeErrorException(s"Call to DeRef but with a non reference type $rtype")}
            }
        }
        
    }
}

def typeOfProgram(p: Program) = p match {
    case TopLevel(e) => {
            val t = typeOf(e, Map())
            println(s"Program type computed successfully as $t")
            t
    }
}
```

Output:
```
defined [32mfunction[39m [36mtypeOf[39m
defined [32mfunction[39m [36mtypeOfProgram[39m
```

<!-- grade_id: cell-bc7a12aa7020a075, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST

/* 
let x : ref(num) = NewRef(10 ) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */

val p1 = Let("x", RefType(NumType), NewRef(Const(10)), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t1 = typeOfProgram(TopLevel(p1))
assert(t1 == NumType, "Test 1 failed: answer should be NumType")
passed(2)
```

Output:
```
Program type computed successfully as NumType

*** Tests Passed (2 points) ***
[36mp1[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m(NumType),
  [33mNewRef[39m([33mConst[39m([32m10.0[39m)),
  [33mLet[39m([32m"dummy"[39m, NumType, [33mAssignRef[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m30.0[39m)), [33mDeRef[39m([33mIdent[39m([32m"x"[39m)))
)
[36mt1[39m: [32mType[39m = NumType
```

<!-- grade_id: cell-b0e79cec0732b13b, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p2 = Let("x", RefType(NumType), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t2 = try {
   typeOfProgram(TopLevel(p2))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

Output:
```
*** Tests Passed (2 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))
[36mp2[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m(NumType),
  [33mNewRef[39m([33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))),
  [33mLet[39m([32m"dummy"[39m, NumType, [33mAssignRef[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m30.0[39m)), [33mDeRef[39m([33mIdent[39m([32m"x"[39m)))
)
[36mt2[39m: [32mAny[39m = [32m"OK -- caught a type error exception: Let binding has type RefType(NumType) whereas it is bound to expression of type RefType(FunType(NumType,NumType))"[39m
```

<!-- grade_id: cell-ce1f9c297983715f, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p3 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t3 = try {
   typeOfProgram(TopLevel(p3))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

Output:
```
Program type computed successfully as FunType(NumType,NumType)
java.lang.AssertionError: assertion failed: The program should not receive a typejava.lang.AssertionError: assertion failed: Please throw TypeErrorException(message) when a type failure occurs
  scala.Predef$.assert(Predef.scala:223)
  ammonite.$sess.cmd23$Helper.liftedTree1$1(cmd23.sc:8)
  ammonite.$sess.cmd23$Helper.<init>(cmd23.sc:3)
  ammonite.$sess.cmd23$.<init>(cmd23.sc:7)
  ammonite.$sess.cmd23$.<clinit>(cmd23.sc:-1)
```

<!-- grade_id: cell-d0af48c3617e0684, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(NewRef(35), 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p4 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(NewRef(Const(35)), Const(30) ), DeRef(Ident("x"))) )
val t4 =  typeOfProgram(TopLevel(p4))
assert(t4 == FunType(NumType, NumType), "Test failed")
passed(2)
//END TEST
```

Output:
```
Program type computed successfully as FunType(NumType,NumType)

*** Tests Passed (2 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))
[36mp4[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m([33mFunType[39m(NumType, NumType)),
  [33mNewRef[39m([33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))),
  [33mLet[39m(
    [32m"dummy"[39m,
    NumType,
    [33mAssignRef[39m([33mNewRef[39m([33mConst[39m([32m35.0[39m)), [33mConst[39m([32m30.0[39m)),
    [33mDeRef[39m([33mIdent[39m([32m"x"[39m))
  )
)
[36mt4[39m: [32mType[39m = [33mFunType[39m(NumType, NumType)
```

<!-- grade_id: cell-69076468ac4208e3, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST

val y = Ident("y")
val x = Ident("x")
val z = Ident("z")
val t1 = RefType(NumType)
val t2 = FunType(t1, NumType)
val t3 = RefType(t2)

val v1 = FunCall(DeRef(y), x)
val v2 = Let("z", RefType(BoolType), NewRef(Geq(DeRef(x), Const(25))), v1)
val fdef = FunDef("z", RefType(NumType), Plus(DeRef(z), DeRef(x)))
val v3 = Let("y", t3, NewRef(fdef), v2)
val v4 = Let("x", t1, NewRef(Const(10)), v3)

val p5 = TopLevel(v4)
assert(typeOfProgram(p5) == NumType, "Test failed")
passed(2)

//END TEST
```

Output:
```
Program type computed successfully as NumType

*** Tests Passed (2 points) ***
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36mz[39m: [32mIdent[39m = [33mIdent[39m([32m"z"[39m)
[36mt1[39m: [32mRefType[39m = [33mRefType[39m(NumType)
[36mt2[39m: [32mFunType[39m = [33mFunType[39m([33mRefType[39m(NumType), NumType)
[36mt3[39m: [32mRefType[39m = [33mRefType[39m([33mFunType[39m([33mRefType[39m(NumType), NumType))
[36mv1[39m: [32mFunCall[39m = [33mFunCall[39m([33mDeRef[39m([33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"x"[39m))
[36mv2[39m: [32mLet[39m = [33mLet[39m(
  [32m"z"[39m,
  [33mRefType[39m(BoolType),
  [33mNewRef[39m([33mGeq[39m([33mDeRef[39m([33mIdent[39m([32m"x"[39m)), [33mConst[39m([32m25.0[39m))),
  [33mFunCall[39m([33mDeRef[39m([33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"x"[39m))
)
[36mfdef[39m: [32mFunDef[39
```

## Problem 2 (10 points): Attempted Type System for Division by Zero

In class, we mentioned that developing a __static type system__ for catching run time errors is much harder. Let us try and see what happens.

Let us modify our type grammar to have two different types of numbers: zero and non-zero.

$$\begin{array}{rcl} 
\mathbf{Type} & \rightarrow & NonZeroNumType \\
& \rightarrow & ZeroNumType \\
\end{array}$$

with the following subset of Lettuce just involving arithmetic and let bindings:

$$\begin{array}{rcl}
\mathbf{Expr} & \rightarrow & Const(\mathbf{Double}) \\
& | & Ident(\mathbf{Identifier})\\
& | & Plus(\mathbf{Expr}, \mathbf{Expr})\\
& | & Div(\mathbf{Expr}, \mathbf{Expr})\\
& | & Let(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}, \mathbf{Expr})\\
\end{array}$$

__Goal:__ Guarantee that a numeric value that receives type _ZeroNumType_ must be $0$ and a numeric value that receives _NonZeroNumType_ must always be non-zero.

__Semantic Rules:__ 

__Constant:__ The new rule for constants are

$$\semRule{f \not= 0}{\typeOf(\texttt{Const(f)}, \alpha) = NonZeroNumType}{const-non-zero}$$

$$\semRule{f = 0}{\typeOf(\texttt{Const(f)}, \alpha) = ZeroNumType}{const-zero}$$

__Identifiers:__ The rule for identifiers remains unchanged (see notes).

__Plus:__ The new rules for plus are

$$\semRule{\typeOf(e1, \alpha) = ZeroNumType,\ \typeOf(e2, \alpha) = ZeroNumType}{\typeOf(\texttt{Plus(e1, e2)}, \alpha) = ZeroNumType}{plus-zero}$$

$$\semRule{\typeOf(e1, \alpha) = NonZeroNumType,\ \typeOf(e2, \alpha) = ZeroNumType}{\typeOf(\texttt{Plus(e1, e2)}, \alpha) = NonZeroNumType}{plus-non-zero-1}$$

$$\semRule{\typeOf(e1, \alpha) = ZeroNumType,\ \typeOf(e2, \alpha) = NonZeroNumType}{\typeOf(\texttt{Plus(e1, e2)}, \alpha) = NonZeroNumType}{plus-non-zero-2}$$

$$\semRule{\typeOf(e1, \alpha) = NonZeroNumType,\ \typeOf(e2, \alpha) = NonZeroNumType}{\typeOf(\texttt{Plus(e1, e2)}, \alpha) = NonZeroNumType}{plus-non-zero-3}$$


__Division:__ The rules for division are

$$\semRule{\typeOf(e1, \alpha) = ZeroNumType,\ \typeOf(e2, \alpha) = NonZeroNumType}{\typeOf(\texttt{Div(e1, e2)}, \alpha) = ZeroNumType}{div-zero}$$

$$\semRule{\typeOf(e2, \alpha) = ZeroNumType}{\typeOf(\texttt{Div(e1, e2)}, \alpha) = \mathbf{typeerror}}{div-error}$$

$$\semRule{\typeOf(e1, \alpha) = NonZeroNumType,\ \typeOf(e2, \alpha) = NonZeroNumType}{\typeOf(\texttt{Div(e1, e2)}, \alpha) = NonZeroNumType}{div-ok}$$

__Let Binding:__ Rules are unchanged (see notes).


We claim that if a Lettuce program is annotated with the types ZeroNumType and NonZeroNumType, and the type annotations are according to the rules above, then any attempt to divide a number by zero will result in __typeerror__ rather than a runtime exception thrown when evaluating the program using an interpreter.

__(A)__ Write a Lettuce program that has the correct type annotations 
(you should use the concrete syntax for your answer) but results in a division by zero that is not 
detected as a __typeerror__.


__(B)__ Point out which rule among the ones provided is incorrect (techically an incorrect type rule is called unsound).


<!-- grade_id: cell-455e0249c2a75e73, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE