# CSCI 3155 Assignment 1 : Scala Programming

This assignment asks you to write scala programs. Fill in the programs below. Use the test cases provided to test them. You are also encouraged to write your own test cases.

Derek Thomas

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage(cond: Boolean, testID: String) = { 
    assert(cond, s"Test $testID Failed"); 
    println(s"Test $testID Passed!")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 1 (5 points)

Write a scala function titled `isPalindrome (s)` that tests whether a given input string `s` is a palindrome.

The following string operations will be helpful.

1. `s.size` gives you the size of the string.
2. `s(i)` gives you the `ith` character for `i = 0, ... , s.size-1`
3. For loops in scala are described here: https://www.tutorialspoint.com/scala/scala_for_loop.htm
   Pay careful attention between the to and until forms of the for loop.

<!-- grade_id: cell-794743da17722228, points: 0, grade: false, solution: true -->
```scala
def isPalindrome(s: String): Boolean = { 
    // YOUR CODE HERE
    for(x <- 0 until s.size/2){
        if(s(x) != s(s.size-x-1))
            return false
    }
    return true
 }
```

Output:
```
defined [32mfunction[39m [36misPalindrome[39m
```

<!-- grade_id: cell-082571ccc4ca97d0, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(isPalindrome("hellolleh"), "Test 1 Failed")
assert(!isPalindrome("bababablacksheep"), "Test 2 Failed")
assert(isPalindrome(""), "Test 3 Failed")
assert(isPalindrome("a"), "Test 4 Failed")
assert(isPalindrome("xxxxxxxx"), "Test 5 Failed")
assert(isPalindrome("ablewasiereisawelba"), "Test 6 Failed")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

## Problem 2 (5 points)

Implement a `isPalindromeRec` function that checks if a given string is a palindrome, but without using mutable `var` or without using `for`/`while` loops.

**Hint:** You should consult Scala documentation on strings to figure out:
1. How do we get first and last characters of a string.
2. How do we get substring between first and last character.
3. How to find out if a string is empty or a single character.
https://www.tutorialspoint.com/scala/scala_strings.htm

<!-- grade_id: cell-80da0f71ccafab19, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def isPalindromeRec(s: String): Boolean = {
    if (s.size <= 1 ||(s.head == s.last && isPalindromeRec(s.substring(1,s.size-1)))) 
        return true
    else
        return false
 }
isPalindromeRec("hellolleh")
```

Output:
```
defined [32mfunction[39m [36misPalindromeRec[39m
[36mres44_1[39m: [32mBoolean[39m = true
```

<!-- grade_id: cell-ef781cc6d55287ba, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST

testWithMessage(isPalindromeRec("hellolleh") == true, "1")
testWithMessage(isPalindromeRec("bababablacksheep") == false, "2")
testWithMessage(isPalindromeRec("") == true, "3")
testWithMessage(isPalindromeRec("a") == true, "4")
testWithMessage(isPalindromeRec("xxxxxxxx") == true, "5")
testWithMessage(isPalindromeRec("ablewasiereisawelba") == true, "6")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!
Test 6 Passed!

*** Tests Passed (5 points) ***
```

## Problem 2 (5 points)

Write a scala function XOR to compute the XOR of two boolean inputs.

<!-- grade_id: cell-096de6a7e20071c9, points: 0, grade: false, solution: true -->
```scala
def computeXOR(x1: Boolean, x2: Boolean): Boolean = {
    // YOUR CODE HERE
    // Im assuming you don't want me to just write this: return x1 ^ x2
    if ((x1 && x2) || (!x1 && !x2))
        return false
    else
        return true
}
```

Output:
```
defined [32mfunction[39m [36mcomputeXOR[39m
```

<!-- grade_id: cell-af1b5ad4b16aa37c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage(!computeXOR(true, true), "1")
testWithMessage(computeXOR(true, false), "2")
testWithMessage(computeXOR(false, true), "3")
testWithMessage(!computeXOR(false, false), "4")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!

*** Tests Passed (5 points) ***
```

## Problem 3 (20 points)

We would like you to implement the methods in the class Rational that is supposed to represent fractions $\frac{n}{d}$, where the numerator and denominator are integers $n, d$. Note that it is important to ensure that $d \not= 0$.

Implement the missing methods below.

You should read this blog post about how to convert between numeric types.

https://alvinalexander.com/scala/how-to-convert-between-numeric-types-in-scala-int-long-float-double


<!-- grade_id: cell-a410d01fa666b077, points: 0, grade: false, solution: true -->
```scala
class Rational(val n: Int, val d: Int) {
    // Implement a pretty printer
    override def toString(): String = {
        s"$n/$d"
    }
    
    def isValid(): Boolean = {
        // check if this is a valid fraction.
        // YOUR CODE HERE
        if (this.d == 0)
            return false
        else
            return true
    }
    def toDouble(): Double = {
        // Convert n/ d to double
        // use toDouble method to convert an integer to double
        // YOUR CODE HERE
        return n.toDouble / d // only one needs to be doulbe, scala will automatically convert the other one
    }
    
    def isInteger(): Boolean = {
        // Is this fraction an integer.
        // A fraction is an integer if n is divisible by d
        // YOUR CODE HERE
        if ((n % d) == 0)
            return true
        else
            return false
    }
    
    def floor(): Double = {
        // return the floor of n/d
        // math.floor(f) takes the floor of a double precision number in scala
        // YOUR CODE HERE
        return math.floor(n.toDouble/d)
    }
    
    
    def rationalEquals(r: Rational): Boolean = {
        // Check if this equals r
        // YOUR CODE HERE
        if ((n.toDouble/d) == (r.n.toDouble/r.d))
            return true
        else return false
    }
    
    
    
    def + (r: Rational) : Rational = {
        // Add two rationals and return a new Rational
        // YOUR CODE HERE
        if (d == r.d)
            return new Rational(n+r.n, d)
        else
            return new Rational(n * r.d + r.n * d, d * r.d)
    }
    
    def - (r: Rational): Rational = {
        // compute this - r and return a new Rational
        // YOUR CODE HERE
        if (d == r.d)
            return new Rational(n-r.n, d)
        else
            return new Rational(n * r.d - r.n * d, d * r.d)
    }
    
    def * (r: Rational): Rational = {
        // multiply two rationals and return a new Rational
        // YOUR CODE HERE
        return new Rational(n * r.n, d * r.d)
    }
    
    def / (r: Rational): Rational = {
        // compute this/r and return a new Rational
        // YOUR CODE HERE
        return new Rational(n * r.d, d * r.n)
    }
    
    // We can use your rationalEquals method to overide the == operator as follows
    // Ignore this method for now -- it will be clearer soon
    override def equals(r: Any): Boolean = r match {
        case rat: Rational => this.rationalEquals(rat) 
        case _ => false
    }
    // END Ignore   
}
```

Output:
```
defined [32mclass[39m [36mRational[39m
```

<!-- grade_id: cell-6275b75f19c16246, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage(new Rational(10, 0).isValid() == false, "1 - isValid")
testWithMessage(new Rational(10, 1).isValid() == true, "2 - isValid")
passed(2)
//END TEST
```

Output:
```
Test 1 - isValid Passed!
Test 2 - isValid Passed!

*** Tests Passed (2 points) ***
```

<!-- grade_id: cell-0dabc332060e86bd, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage(new Rational(10, 3).toDouble() == 10.0/3.0, "1 - toDouble")
testWithMessage(new Rational(19, 5).toDouble() == 19.0/5.0, "2 - toDouble")
passed(2)
//END TEST
```

Output:
```
Test 1 - toDouble Passed!
Test 2 - toDouble Passed!

*** Tests Passed (2 points) ***
```

<!-- grade_id: cell-cd465592f873c02f, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage(new Rational(15, 3).isInteger() == true, "1 - isInteger")
testWithMessage(new Rational(15, 8).isInteger() == false, "2 - isInteger")
testWithMessage(new Rational(225, 15).isInteger() == true, "3 - isInteger")
testWithMessage(new Rational(181, 2).isInteger() == false, "4 - isInteger")
passed(2)
//END TEST
```

Output:
```
Test 1 - isInteger Passed!
Test 2 - isInteger Passed!
Test 3 - isInteger Passed!
Test 4 - isInteger Passed!

*** Tests Passed (2 points) ***
```

<!-- grade_id: cell-1746cc35b42175e3, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage(new Rational(15, 4).floor() == 3.0, "1 - floor")
testWithMessage(new Rational(100, 45).floor() == 2.0, "2 - floor")
testWithMessage(new Rational(-25, 3).floor() == -9.0, "3 - floor")
testWithMessage(new Rational(-40, 7).floor() == -6.0, "4 - floor")
passed(2)
//END TEST
```

Output:
```
Test 1 - floor Passed!
Test 2 - floor Passed!
Test 3 - floor Passed!
Test 4 - floor Passed!

*** Tests Passed (2 points) ***
```

<!-- grade_id: cell-1c01befb23ee7487, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage((new Rational(20, 5)) == (new Rational(16 , 4)), "1 - rationalEquals")
testWithMessage((new Rational(24, 5)) == (new Rational(120, 25)), "2 - rationalEquals")
testWithMessage((new Rational(19, 5)) != (new Rational(58, 15)), "3 - rationalEquals")
testWithMessage((new Rational(24, 1)) != (new Rational(120, 15)), "4 - rationalEquals")
passed(2)
//END TEST
```

Output:
```
Test 1 - rationalEquals Passed!
Test 2 - rationalEquals Passed!
Test 3 - rationalEquals Passed!
Test 4 - rationalEquals Passed!

*** Tests Passed (2 points) ***
```

<!-- grade_id: cell-f595d18030ee424d, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val r1 = new Rational(10, 3)
val r2 = new Rational(14, 3)
val r3 = r1 + r2
testWithMessage(r3.n == r3.d * 8, "1 - addition of 10/3 + 14/3 - ")

val r4 = new Rational(9, 2)
val r5 = new Rational(18, 2)
val r6 = r4 + r5
testWithMessage(r6.n * 2 == r6.d * 27, "2 - addition of 9/2 + 18/2 - ")

val r7 = new Rational(15, 7)
val r8 = new Rational(19, 6)
val r9 = r7 + r8
testWithMessage(r9.n * 42 == r9.d * 223, "3 - addition of 15/7 + 19/6 - ")

val r10 = r1 - r2
testWithMessage(r10.n * 3 == r10.d * -4, "1 - subtraction of 10/3 - 14/3 - ")

val r11 = r6 - r7
passed(10)
//END TEST
```

Output:
```
Test 1 - addition of 10/3 + 14/3 -  Passed!
Test 2 - addition of 9/2 + 18/2 -  Passed!
Test 3 - addition of 15/7 + 19/6 -  Passed!
Test 1 - subtraction of 10/3 - 14/3 -  Passed!

*** Tests Passed (10 points) ***
[36mr1[39m: [32mRational[39m = 10/3
[36mr2[39m: [32mRational[39m = 14/3
[36mr3[39m: [32mRational[39m = 24/3
[36mr4[39m: [32mRational[39m = 9/2
[36mr5[39m: [32mRational[39m = 18/2
[36mr6[39m: [32mRational[39m = 27/2
[36mr7[39m: [32mRational[39m = 15/7
[36mr8[39m: [32mRational[39m = 19/6
[36mr9[39m: [32mRational[39m = 223/42
[36mr10[39m: [32mRational[39m = -4/3
[36mr11[39m: [32mRational[39m = 159/14
```

## Problem 4 (20 points)

Newton invented the Newton-Raphson method for solving an equation. We are going to ask you to write some code to solve equations.

To solve an equation of the form

$$ x^2 - 3x + 2 = 0$$

we start from an initial guess at the solution: say $x_0 = 4.5$
Each time we have the $i^{th}$ guess $x_i$, we update it as

$$ x_{i+1} = x_i - \frac{f(x_i)}{f'(x_i)} $$

For our equation, $f(x) = x^2 - 3x +2$ and $f'(x) = 2 x - 3$. Thus, our update equation is 
$$ x_{i+1} = x_i - \frac{x_i^2 - 3 x_i + 2}{2 x_i - 3}$$.

We stop whenever $|f(x_i)| \leq 10^{-8}$: i.e, we are very close to a root of the function. Gory details are here for the curious ones (curiosity = bad for cats, great for students).
http://www.math.ubc.ca/~anstee/math104/newtonmethod.pdf

__(A, 5 points)__ Write a scala function `calculateRoot(x0: Double)` that takes in the initial guess $x_0$ as input and calculates the root for the polynomial $x^2 - 3 x + 2$. Your code should use a While loop.

Caution: To calculate $x^2$, you should use `math.pow(x,2)` in Scala.

<!-- grade_id: cell-2dd21341b51d86ad, points: 0, grade: false, solution: true -->
```scala
def calculateRoot(x0: Double): Double = { 
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    def df(x: Double) = { 2* x - 3}
    // Your code here: use f(x) and df(x) provided already.
    // YOUR CODE HERE
    var xi = x0
    while (math.abs(f(xi)) > math.pow(10,-8))
        xi = xi - f(xi) / df(xi)
    return xi
}
```

Output:
```
defined [32mfunction[39m [36mcalculateRoot[39m
```

<!-- grade_id: cell-2e7493c157d154dc, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
def isRootCorrect(x: Double): Boolean = {
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    val y = calculateRoot(x)
    math.abs(f(y)) <= 1e-08
}

testWithMessage(isRootCorrect(3.5), "1")
testWithMessage(isRootCorrect(4.6), "2")
testWithMessage(isRootCorrect(-3), "3")
testWithMessage(isRootCorrect(-3.5), "4")
testWithMessage(isRootCorrect(1.3), "5")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
defined [32mfunction[39m [36misRootCorrect[39m
```

__(B, 5 points)__ Implement the function `calculateRootRec` that uses a recursion instead of a while loop to compute the root. You should __not__ use a mutable `var` in your code. 

<!-- grade_id: cell-57d2ad2f5ba44329, points: 0, grade: false, solution: true -->
```scala
def calculateRootRec(x: Double): Double = {
    // YOUR CODE HERE
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    def df(x: Double) = { 2* x - 3}
    if (math.abs(f(x)) > math.pow(10,-8))
        return calculateRoot(x - f(x) / df(x))
    else 
        return x
}
```

Output:
```
defined [32mfunction[39m [36mcalculateRootRec[39m
```

<!-- grade_id: cell-49bff1b838e3e910, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
def isRootCorrectRec(x: Double): Boolean = {
    val y = calculateRootRec(x)
    (math.abs(y - 2.0) <= 1e-06) || (math.abs(y - 1.0) <= 1e-06)
}

testWithMessage(isRootCorrectRec(3.5), "1")
testWithMessage(isRootCorrectRec(4.6), "2")
testWithMessage(isRootCorrectRec(-3), "3")
testWithMessage(isRootCorrectRec(-3.5), "4")
testWithMessage(isRootCorrectRec(1.3), "5")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
defined [32mfunction[39m [36misRootCorrectRec[39m
```

__(C, 10 points)__ Write a generalized Newton-Raphson function `solveEquationNewtonRaphson` to solve any 
equation $f(x) = 0$ wherein the functions $f$ and $f'(x)$ are given as inputs. Do not use while loops or mutable `var` in your code.



<!-- grade_id: cell-d053a7e0f075bd32, points: 0, grade: false, solution: true -->
```scala
def solveEquationNewtonRaphson(x: Double, f: Double => Double, df: Double => Double): Double = {
   // YOUR CODE HERE
   if (math.abs(f(x)) > math.pow(10,-8))
        return solveEquationNewtonRaphson(x - f(x) / df(x), f, df)
    else 
        return x
}
```

Output:
```
defined [32mfunction[39m [36msolveEquationNewtonRaphson[39m
```

<!-- grade_id: cell-4a86f811b0471af9, points: 2, grade: true, solution: false -->
```scala
// Test Cases
def f1(x: Double) = 2.0 * math.pow(x,3) - math.pow(x, 2) - 1.0
def df1(x: Double) = 6.0 * math.pow(x, 2) - 2 * x 
val x1 = solveEquationNewtonRaphson(3.0, f1, df1)
testWithMessage(math.abs(f1(x1)) <= 1e-06, "1")
passed(2)
```

Output:
```
Test 1 Passed!

*** Tests Passed (2 points) ***
defined [32mfunction[39m [36mf1[39m
defined [32mfunction[39m [36mdf1[39m
[36mx1[39m: [32mDouble[39m = [32m1.000000000001152[39m
```

<!-- grade_id: cell-dd5de96626a8f24c, points: 2, grade: true, solution: false -->
```scala
def f2(x: Double) = math.sin(x) - 0.5 * x 
def df2(x: Double) = math.cos(x) - 0.5
val x2 = solveEquationNewtonRaphson(0.7, f2, df2)
testWithMessage(math.abs(f2(x2)) <= 1e-06, "2")
passed(2)
```

Output:
```
Test 2 Passed!

*** Tests Passed (2 points) ***
defined [32mfunction[39m [36mf2[39m
defined [32mfunction[39m [36mdf2[39m
[36mx2[39m: [32mDouble[39m = [32m8.537195744177245E-13[39m
```

<!-- grade_id: cell-24787c8854b20def, points: 2, grade: true, solution: false -->
```scala
def f3(x: Double) = math.log(x) - 2 
def df3(x: Double) = 1/x 
val x3 = solveEquationNewtonRaphson(2.0, f3, df3)
testWithMessage(math.abs(f3(x3)) <= 1e-06, "3")
passed(2)
```

Output:
```
Test 3 Passed!

*** Tests Passed (2 points) ***
defined [32mfunction[39m [36mf3[39m
defined [32mfunction[39m [36mdf3[39m
[36mx3[39m: [32mDouble[39m = [32m7.389056098804318[39m
```

<!-- grade_id: cell-08b0bc89a3d1d886, points: 2, grade: true, solution: false -->
```scala
def f4(x: Double) = math.atan(x) - 1
def df4(x: Double) = 1.0/(1.0 + x*x)
val x4 = solveEquationNewtonRaphson(3.0, f4, df4)
testWithMessage(math.abs(f4(x4)) <= 1e-06, "4")
passed(2)
```

Output:
```
Test 4 Passed!

*** Tests Passed (2 points) ***
defined [32mfunction[39m [36mf4[39m
defined [32mfunction[39m [36mdf4[39m
[36mx4[39m: [32mDouble[39m = [32m1.5574077246540006[39m
```

<!-- grade_id: cell-e858482df5178f7d, points: 2, grade: true, solution: false -->
```scala
def f5(x: Double) = math.cos(x)* math.cos(x) - 1
def df5(x: Double) = -2.0 * math.sin(x) * math.cos(x)
val x5 = solveEquationNewtonRaphson(3.5, f5, df5)
testWithMessage(math.abs(f5(x5)) <= 1e-06, "5")
passed(2)
```

Output:
```
Test 5 Passed!

*** Tests Passed (2 points) ***
defined [32mfunction[39m [36mf5[39m
defined [32mfunction[39m [36mdf5[39m
[36mx5[39m: [32mDouble[39m = [32m3.1416751165142625[39m
```

-- That is all for now, folks! -- 