# CSCI 3155 Recitation 8
October 19, 2018

## Continuations
Continuations are functions we pass as arguments to other functions which capture what to "do next". Instead of writing:
```
def func() {
    val x = doThing1()
    val y = doThing2(x, 3)
    doThing3()
}
```
we can make every function call recursive by changing it to continuation passing style:
```
def func[X, Y](k: X => Y) {
    doThing1(x => DoThing2(x, 3, y => doThing3(k))
}

// Same as:
def func[X, Y](k: X => Y) {
    doThing1(          x =>
    doThing2(x, 3,     y =>
    doThing3(          k))
}
```

### Examples

https://docs.google.com/presentation/d/1OlAu7pjBfKogQg6Z5Ykkid9FfWVH-yrSbPUCHMI-3HM/edit?usp=sharing

### Exercise: Backtracking
As with any good programing tool, we can use continuations to solve problems without putting them everywhere in our code. In this example, you will write a search function that looks for a value in a binary tree, **not** a B**S**T. Take advantage of continuations to remember where to look if something isn't found while keeping the function tail recursive.

<!-- grade_id: cell-bdaf875cad5b0678, points: 0, grade: false, solution: true -->
```scala
sealed trait Tree
case object Empty extends Tree
case class Node(l: Tree, d: Int, r: Tree) extends Tree

def search(t: Tree, i: Int, fail_continuation: () => Boolean): Boolean =
    // BEGIN SOLUTION
    t match {
        case Empty => fail_continuation()
        case Node(l, j, r) if i == j =>
            true
        case Node(l, j, r) =>
            search(l, i, () => {
                search(r, i, () => {
                    fail_continuation()
                })
            })
    }
    //END SOLUTION

```

Output:
```
defined [32mtrait[39m [36mTree[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mNode[39m
defined [32mfunction[39m [36msearch[39m
```

<!-- grade_id: cell-df2188387bae0d26, points: 1, grade: true, solution: false -->
```scala
val t = Node(Empty, 10, Node(Empty, 6, Empty))
assert(search(t, 10, () => false))
assert(!search(t, 0, () => false))
```

Output:
```
[36mt[39m: [32mNode[39m = Node(Empty,10,Node(Empty,6,Empty))
```

### Exercise: Eval (again)

This is similar to the example in class: implement eval for the small language given below, ensuring that every function call is a tail call.

<!-- grade_id: cell-83c17e7eeba4b4e4, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class BoolLiteral(b: Boolean) extends Expr
case class And(left: Expr, right: Expr) extends Expr
case class If(test: Expr, then: Expr, otherwise: Expr) extends Expr

def eval(e: Expr, continuation: Boolean => Boolean): Boolean =
    // BEGIN SOLUTION
    e match {
        case BoolLiteral(b) =>
            continuation(b)
        case And(left, right) =>
            eval(left, left_val =>
                 eval(right, right_val =>
                      continuation(left_val && right_val)
                     )
                )
        case If(test, then, otherwise) =>
            eval(test, test_val =>
                 if (test_val) eval(then, b => continuation(b))
                 else eval(otherwise, continuation)
                )
    }
    // END SOLUTION
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mBoolLiteral[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mIf[39m
defined [32mfunction[39m [36meval[39m
```

<!-- grade_id: cell-c127bcc04c6197cc, points: 1, grade: true, solution: false -->
```scala
val e = If(
    And(BoolLiteral(true), BoolLiteral(true)),
    BoolLiteral(false),
    BoolLiteral(true)
)
assert(!eval(e, x => x))
```

Output:
```
[36me[39m: [32mIf[39m = If(And(BoolLiteral(true),BoolLiteral(true)),BoolLiteral(false),BoolLiteral(true))
```

## Polymorphism

Polymorphism allows us to write reusable, type safe code. This is also called type parameters, generics, template parameters, or type variables.

As a reminder, we list type parameters in brackets similarly to how we list regular parameters in parentheses:

```scala
def id[T](t: T): T = t
// OR
case class C[T](t: T) {
    def get: T = t
}
```

Output:
```
defined [32mfunction[39m [36mid[39m
defined [32mclass[39m [36mC[39m
```

Type inference works just the same for these as it does for concrete types, and the scope of type parameters is the function body, like regular parameters:

```scala
def f[T](t: T): T = {
    val t2: T = t
    val t3    = t2
    t3
}
```

Output:
```
defined [32mfunction[39m [36mf[39m
```

### Exercise: Reusable data structures
Make a reusable tree data structure which can hold any arbitrary type at it's nodes, as long as all nodes in the tree have the same type. Use the constructors `Empty` and `Node`

<!-- grade_id: cell-89ce853776b2aa84, points: 0, grade: false, solution: true -->
```scala
sealed trait TreePoly[T]
// BEGIN SOLUTION
case class EmptyPoly[T]() extends TreePoly[T]
case class NodePoly[T](l: TreePoly[T], d: T, r: TreePoly[T]) extends TreePoly[T]
// END SOLUTION
```

Output:
```
defined [32mtrait[39m [36mTreePoly[39m
defined [32mclass[39m [36mEmptyPoly[39m
defined [32mclass[39m [36mNodePoly[39m
```

```scala
// This should compile:
val t = NodePoly(
    NodePoly(
        EmptyPoly(),
        5,
        EmptyPoly()),
    6,
    EmptyPoly())
```

Output:
```
[36mt[39m: [32mNodePoly[39m[[32mInt[39m] = NodePoly(NodePoly(EmptyPoly(),5,EmptyPoly()),6,EmptyPoly())
```

```scala
// This should NOT compile:
NodePoly(
    NodePoly(
        EmptyPoly(),
        5,
        EmptyPoly()),
    "6",
    EmptyPoly())
```

Output:
```
cmd9.sc:2: type mismatch;
 found   : cmd9Wrapper.this.cmd7.wrapper.NodePoly[Int]
 required: cmd9Wrapper.this.cmd7.wrapper.TreePoly[Any]
Note: Int <: Any, but trait TreePoly is invariant in type T.
You may wish to define T as +T instead. (SLS 4.5)
    NodePoly(
            ^Compilation Failed
```

> _Side-note_: for those that are anoyed by the need to make a "new" empty node each time as opposed to before where we could just have one `case object`, we can get around it as follows (see https://stackoverflow.com/questions/7399044/scala-upper-type-bounds-and-parent-classes)
> 
> ```
sealed trait Tree[+T]
case object Empty extends Tree[Nothing]
```

### Exercise: generic functions
Implement map, which should allow someone to take a `Tree[X]` and a function `X => Y` and get out a `Tree[Y]`. Do this with continuations.

<!-- grade_id: cell-11843a63b5b7990f, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def map[X, Y](t: TreePoly[X], f: X => Y, continuation: TreePoly[Y] => TreePoly[Y]): TreePoly[Y] = t match {
    case EmptyPoly() =>
        continuation(EmptyPoly())
    case NodePoly(l, d, r) =>
        map(l, f, l_mapped =>
            map(r, f, r_mapped =>
                continuation(NodePoly(l_mapped, f(d), r_mapped))))
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mmap[39m
```

<!-- grade_id: cell-d725bf4fbe8683ac, points: 1, grade: true, solution: false -->
```scala
// TEST
val t2 = NodePoly(
    NodePoly(
        EmptyPoly(),
        "5",
        EmptyPoly()),
    "6",
    EmptyPoly())

assert(map(t, (i: Int) => i.toString, id[TreePoly[String]]) == t2)
// TEST
```

Output:
```
[36mt2[39m: [32mNodePoly[39m[[32mString[39m] = NodePoly(NodePoly(EmptyPoly(),5,EmptyPoly()),6,EmptyPoly())
```